<a href="https://colab.research.google.com/github/notpromising/Python-Assignments/blob/main/Graphs_Python.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Parth Mahale"
COLLABORATORS = ""
```

---

# Assignment 10: Graphs

_This notebook is a derivative work of ["Expressions as Classes"](https://colab.research.google.com/drive/1qO-GOb5KJz9WucdOswISkM8hILwLDUo9) by [Luca de Alfaro](https://sites.google.com/a/ucsc.edu/luca/home), which is licensed under [CC BY-NC-ND 4.0](https://creativecommons.org/licenses/by-nc-nd/4.0/), distributed by permission of the original author.  This notebook is licensed under [CC BY-NC-ND 4.0](https://creativecommons.org/licenses/by-nc-nd/4.0/) by [Peter Alvaro](https://users.soe.ucsc.edu/~palvaro/)._

# Instructions

## Notebook Format

This is a homework notebook.  It consists of various types of cells: 

* Text cells: you should read them!
* Code cells: you should run them, as they may set up the problems that you are asked to solve.
* **Solution** cells: These are cells where you should enter a solution.  You will see a marker in these cells that indicates where your work should be inserted:  

```
    # YOUR CODE HERE
```    

* Test cells: These cells contain some tests, and are worth some points.  You should run the cells as a way to debug your code, and to see if you understood the question, and whether the output of your code is produced in the correct format.

**When we grade your notebook, we will run many additional tests in addition to the ones you see here.  (Otherwise, you'd be able to get full credit by hard-coding the desired output!)**

## Working on Your Notebook

To work on your notebook, we recommend using Colab.  Working in Colab has many benefits:

  * You don't have to maintain a working Python environment on your own machine; you can work from any machine that has an internet connection.
  * Colab preserves the revision history, which is useful for many reasons.
  * Your work is automatically saved in Google Drive.
  
In Colab, go to "File > Save a copy in Drive..." and then you'll have your own copy to work on.

## Submitting Your Notebook

Before you turn your finished notebook in, it's a good idea to sure everything runs as expected. First, **factory reset the runtime** in Colab (go to "Runtime > Factory reset runtime").  Then, **run all the cells** (go to "Runtime > Run all") in Colab.

Submit your work as follows: 

  * Download the notebook from Colab, clicking on "File > Download .ipynb".
  * Upload the resulting file to [this Google form](https://docs.google.com/forms/d/e/1FAIpQLSdZGL7tDOWYSaiTOQUz9snLSxdzjO3Q4DWkAjweFgcVvaxvVA/viewform?usp=sf_link).
  * **Deadline: 9pm Tuesday, March 3.**

You can submit multiple times, and the last submission before the deadline will be used to assign you a grade.

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and the names of anyone you collaborated with in the below cells.  For collaborators, list anyone who (for example) discussed the general techniques involved with you, or pointed you to useful Python documentation -- collaboration of that kind is encouraged.  Remember that sharing code or using others' code is not allowed.

```python
NAME = "Parth Mahale"
COLLABORATORS = ""
```

## What Happens Next? 

After you submit your notebook, your instructor at some point will retreat to a secret hideout, put on some good music, and run some mysterious scripts.  These will generate two things: 

  * Your grade, which goes into a spreadsheet. 
  * Feedback, shared to you as a PDF file on Google Drive.  The PDF shows your work, your grade, the tests that passed and those that failed, and any comments left by the instructor and the TAs.

# Testing

Make sure to run the following cell to make sure that the Python testing framework, `nose`, is installed.  (It's required for the rest of the notebook to work.)

```python
try:
    from nose.tools import assert_equal, assert_almost_equal
    from nose.tools import assert_true, assert_false
    from nose.tools import assert_not_equal, assert_greater_equal
except:
    !pip install nose
    from nose.tools import assert_equal, assert_almost_equal
    from nose.tools import assert_true, assert_false
    from nose.tools import assert_not_equal, assert_greater_equal
```

## Graphs

In lecture Tuesday, we covered graphs in the abstract, and then went on to implement a graph class endowed with methods to add nodes and edges, to determine the successors of a given node, and ultimately to determine the set of nodes that are *reachable* from a given starting node (the transitive closure of successor).

So far, all of the graph operations we studied involve a single graph.  Often, we want to combine or compare multiple graphs, much as we do with collections.

What are useful, general operations over pairs of graphs that we may implement?  Here are a few. 

*   **Union** and **intersection**. 
*   **Induced:** given a graph $G = (V, E)$ and a set of vertices $U$, we return the graph with set of vertices $V \cap U$ and set of edges $E \cap (V\cap U \times V \cap U)$. This is the portion of the original graph that only involves vertices in V. 
* **Difference:** Remove, from a graph $G$, all vertices in a specified set $U$, along with the edges that have an endpoint in $U$. 

Let us add these implementations.

```python
import networkx as nx # Library for displaying graphs.

class Graph(object):
    
    def __init__(self, vertices=None, edges=None):
        self.s = {u: set() for u in vertices or []}
        for u, v in (edges or []):
            self.add_edge((u, v))
        
    def show(self):
        g = nx.DiGraph()
        g.add_nodes_from(self.s.keys())
        g.add_edges_from([(u, v) for u in self.s for v in self.s[u]]) 
        nx.draw(g, with_labels=True)
        

    @property
    def vertices(self):
        return set(self.s.keys())

    @property
    def edges(self):
        return {(u, v) for u, d in self.s.items() for v in d}

    def __eq__(self, other):
        """We need to define graph equality."""
        if self.vertices != other.vertices:
            return False
        for v, d in self.s.items():
            if d != other.s[v]:
                return False
        return True
    
    def __repr__(self):
        r = "Graph:"
        for v in self.vertices:
            r += "\n %r : %r" % (v, self.s.get(v))
        return r
                    
    def show(self):
        g = nx.DiGraph()
        g.add_nodes_from(self.vertices)
        g.add_edges_from([(u, v) for u in self.vertices for v in self.s[u]]) 
        nx.draw(g, with_labels=True)
        
    def add_vertex(self, v):
        self.vertices.add(v)
        # We must be careful not to overwrite the successor relation
        # in case v might already be present in the graph.
        self.s[v] = self.s.get(v, set())
        
    def add_edge(self, e):
        """Adds an edge e = (u, v) between two vertices u, v.  If the 
        two vertices are not already in the graph, adds them."""
        u, v = e
        self.vertices.update({u, v})
        # Initializes the successor function if needed.
        self.s[u] = self.s.get(u, set()) | {v}
        self.s[v] = self.s.get(v, set())
        
    def successors(self, u):
        """Returns the set of successors of a vertex u"""
        return self.s[u]
    
    def reachable(self, v):
        """Returns the set of vertices reachable from an initial vertex v."""
        vopen = {v}
        vclosed = set()
        while len(vopen) > 0:
            u = vopen.pop()
            vclosed.add(u)
            vopen.update(self.s[u] - vclosed)
        return vclosed
    
    def __and__(self, g):
        """Returns the intersection of the current graph with a 
        specified graph g."""
        return Graph(vertices=self.vertices & g.vertices, 
                     edges=self.edges & g.edges)
    
    def induced(self, vertex_set):
        """Returns the subgraph induced by the set of vertices vertex_set."""
        common_vertices = vertex_set & self.vertices 
        gg = Graph(vertices = common_vertices)
        for v in common_vertices:
            gg.s[v] = self.s[v] & common_vertices
        gg._check()
        return gg


```

## Problem 1: Implement Graph Union


Write a `union` method for a graph so that, for $G_1 = (V_1, E_1)$ and $G_2 = (V_2, E_2)$, with $G_1$ represented by `g1` in code and $G_2$ represented by `g2`, 

    g1 | g2

returns the graph $G_{12} = (V_1 \cup V_2, E_1 \cup E_2)$ having as vertices the union of the vertices of $G_1$ and $G_2$, and as edges the union of the edges of $G_1$ and $G_2$.


<!-- grade_id: cell-612487afbefac1d2, points: 0, grade: false, solution: true -->
```python
# Implement graph union

def graph_union(self, g):
    """Returns the union of the current graph, and of the graph g."""
    u_edge=self.edges.union(g.edges)
    u_vert=self.vertices.union(g.vertices)
    g= Graph(u_vert, u_edge)
    
    return g

Graph.__or__ = graph_union
```

<!-- grade_id: cell-67f93d6ff88cb460, points: 10, grade: true, solution: false -->
```python
#Tests for graph union 

# Disjoint graphs. 
g1 = Graph(vertices=[1, 3, 2], edges=[(1, 3), (2, 3)])
g2 = Graph(vertices=[4, 5], edges=[(4, 5)])
g = g1 | g2
assert_equal(g.vertices, {1, 2, 3, 4, 5})
assert_true((2, 3) in g.edges)
assert_true((4, 5) in g.edges)
assert_false((1, 4) in g.edges)
g3 = g2 | g1
assert_equal(g, g3)
```

<!-- grade_id: cell-956350eb04c6bcec, points: 10, grade: true, solution: false -->
```python
# More tests for graph union

# Overlapping graphs.
g1 = Graph(vertices=['a', 'b', 'c'], edges=[('a', 'b'), ('b', 'c')])
g2 = Graph(vertices=['b', 'c', 'd', 'e'], edges=[('c', 'd'), ('b', 'c')])
g = g1 | g2
assert_equal(g.vertices, {'a', 'b', 'c', 'd', 'e'})
assert_equal(g.edges, {('a', 'b'), ('b', 'c'), ('c', 'd')})
g3 = g2 | g1
assert_equal(g, g3)
```

<!-- grade_id: cell-e7a96367aee23c84, points: 10, grade: true, solution: false -->
```python
# Even more tests for graph union

# Empty graph.
g1 = Graph(vertices=[1, 3, 2], edges=[(1, 3), (2, 3)])
g2 = Graph()
g = g1 | g2
assert_equal(g, g1)
```

## Problem 2:  Is a graph a tree? 

A tree is a graph $(V, E)$ with two special properties: 

* Every vertex has at most one incoming edge (or a single *parent*, as we discussed in class).
* Either there are no vertices, or there is a vertex with no incoming edges (parents), called the _root_, from which all other vertices are reachable. 

If the second property does not hold, incidentally, the graph is called a _forest._

Write an `is_tree` property that has value True if the graph is a tree, and has value False otherwise.

<!-- grade_id: cell-cc5d32aa6cd9babe, points: 0, grade: false, solution: true -->
```python
 #Implementation of tree test

def graph_is_tree(self):
    """Returns True iff the graph is a tree."""
    if type(self) != Graph: 
      return False
    for i in self.vertices:
      if len(self.reachable(i)) < 1:
        return True
      

    n = None
    for i, j in g.edges:
      if j == n: 
        return False
      n = j
    return True

Graph.is_tree = property(graph_is_tree)
```

<!-- grade_id: cell-06ff4eefde003bcf, points: 10, grade: true, solution: false -->
```python
 #Tests for `is_tree`

g = Graph(vertices=[1, 2, 3], edges=[(1, 2), (1, 3)])
assert_true(g.is_tree)
g = Graph(vertices=[1, 2, 3], edges=[(1, 2), (2, 3), (1, 3)])
assert_false(g.is_tree)
g = Graph(vertices=[1, 2, 3], edges=[(1, 3), (2, 3)])
assert_false(g.is_tree)
```

<!-- grade_id: cell-3e480dc41ae6cd4f, points: 10, grade: true, solution: false -->
```python
#More tests for `is_tree`

g = Graph()
assert_true(g.is_tree)

```